---
title: 'Chapter 2'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 9715ec9301
xp: 100
```



`@instructions`
Stepping through the code, we will start by loading the appropriate libraries.

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
